---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 8680baadc8
xp: 50
```

`@projector_key`
c3c8edd14ef3e3564996c0edd15910c4

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 9c0110c3e1
xp: 50
```

`@projector_key`
7537fb354f51559aa9b6c8426de99414
